<a href="https://colab.research.google.com/github/Nikunj-adsum24/personalce800/blob/main/week10/problem1.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "NIKUNJKUMAR DHANANI"
COLLABORATORS = "2111709"
```

---

---
# Welcome to CE880
### This is your week-10 : Problem notebook

For this problem set, we'll be using the Jupyter notebook and please upload this notebook to [Google Colab](https://colab.research.google.com/). 



```python
import numpy as np
from sklearn import datasets, linear_model
from sklearn.metrics import mean_squared_error, r2_score

# Load the diabetes dataset
diabetes_X, diabetes_y = datasets.load_diabetes(return_X_y=True)

# Use only one feature
diabetes_X = diabetes_X[:, np.newaxis, 2]

# Split the data into training/testing sets
diabetes_X_train = diabetes_X[:-20]
diabetes_X_test = diabetes_X[-20:]

# Split the targets into training/testing sets
diabetes_y_train = diabetes_y[:-20]
diabetes_y_test = diabetes_y[-20:]
```

---
## Question 1: 

#### Regression
Use the diabetes dataset from sklearn, which I have spilted into training and test sets (diabetes_X_train, diabetes_y_train,diabetes_X_test,diabetes_y_test), please train the `Ridge` linear regression model to get the test accuracy and return `mean squared error`.  

<!-- grade_id: my_reg_model_mse_ans, points: 0, grade: false, solution: true -->
```python
def my_reg_model_mse(diabetes_X_train, diabetes_y_train,diabetes_X_test,diabetes_y_test):
    """Write a function to train a machine learning model to get the `mean_squared_error` on diabetes_X_test. 
    Please use Ridge() model to train a regressor. """
    
    # YOUR CODE HERE
    from sklearn.svm import SVR
    from sklearn.metrics import mean_squared_error
    svr = SVR(kernel = 'linear', C = 250)
    svr.fit(diabetes_X_train, diabetes_y_train)
    y_test_pred = svr.predict(diabetes_X_test)
    mse = mean_squared_error(diabetes_y_test, y_test_pred)
    print(mse)
    return mse
    raise NotImplementedError()
```

<!-- grade_id: my_reg_model_mse_test, points: 5, grade: true, solution: false -->
```python
# Check you solution by running this cell
import math
assert math.isclose(my_reg_model_mse(diabetes_X_train, diabetes_y_train,diabetes_X_test,diabetes_y_test), 
                    3604.26, rel_tol = 0.05)
```

Output:
```
2551.2549147343525
AssertionError                            Traceback (most recent call last)
<ipython-input-10-8e05269cb419> in <module>()
      2 import math
      3 assert math.isclose(my_reg_model_mse(diabetes_X_train, diabetes_y_train,diabetes_X_test,diabetes_y_test), 
----> 4                     3604.26, rel_tol = 0.05)

AssertionError:
```

---
## Question 2: 

#### Regression
Use the diabetes dataset from sklearn, which I have spilted into training and test sets (diabetes_X_train, diabetes_y_train,diabetes_X_test,diabetes_y_test), please train the `Ridge` linear regression model to get the test accuracy and return `r2` (r sqaured value).  

<!-- grade_id: my_reg_model_r2_ans, points: 0, grade: false, solution: true -->
```python
def my_reg_model_r2(diabetes_X_train, diabetes_y_train,diabetes_X_test,diabetes_y_test):
    """Write a function to train a machine learning model to get the `r2_score` on diabetes_X_test data . 
    Please use LinearRegression model to train a regressor. """
    # YOUR CODE HERE
    from sklearn.svm import SVR
    from sklearn.metrics import r2_score
    svr = SVR(kernel = 'linear', C = 250)
    svr.fit(diabetes_X_train, diabetes_y_train)
    y_test_pred = svr.predict(diabetes_X_test)
    r2_score = r2_score(diabetes_y_test, y_test_pred)
    print(r2_score)
    return r2_score
    raise NotImplementedError()
```

<!-- grade_id: my_reg_model_r2_test, points: 4, grade: true, solution: false -->
```python
# Check you solution by running this cell
import math
assert math.isclose(my_reg_model_r2(diabetes_X_train, diabetes_y_train,diabetes_X_test,
                                    diabetes_y_test), 0.2539, rel_tol = 0.05)
```

Output:
```
0.4719167001849758
AssertionError                            Traceback (most recent call last)
<ipython-input-11-ef50329c6566> in <module>()
      2 import math
      3 assert math.isclose(my_reg_model_r2(diabetes_X_train, diabetes_y_train,diabetes_X_test,
----> 4                                     diabetes_y_test), 0.2539, rel_tol = 0.05)

AssertionError:
```